---
title: 'Analyzing Twitter Data in R'
description: 'By the end of the chapter, the students will have understood the importance of social media analysis in R and, in particular, some of the benefits from analyzing Twitter data. They would have written R code to connect to live Twitter data, explore their structure, perform some data cleaning operations, sentiment analysis and visualization on a leaflet map.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Retrieving and manipulating Twitter data

```yaml
type: VideoExercise
key: 7adb3cc3ab
xp: 50
```

`@projector_key`
c4a6a7716624517ffa84d7c94c086e39
